# Multiple dispatch

In this notebook we'll explore **multiple dispatch**, which is a key feature of Julia.

Multiple dispatch makes software *generic* and *fast*!

#### Starting with the familiar

To understand multiple dispatch in Julia, let's start with what we've already seen.

We can declare functions in Julia without giving Julia any information about the types of the input arguments that function will receive:

```julia
f(x) = x^2
```

Output:
```
f (generic function with 1 method)
```

and then Julia will determine on its own which input argument types make sense and which do not:

```julia
f(10)
```

Output:
```
100
```

```julia
f([1, 2, 3])
```

Output:
```
[2] literal_pow at ./none:0 [inlined]
 [3] f(::Array{Int64,1}) at ./In[1]:1
 [4] top-level scope at In[3]:1
```

#### Specifying the types of our input arguments

However, we also have the *option* to tell Julia explicitly what types our input arguments are allowed to have.

For example, let's write a function `foo` that only takes strings as inputs.

```julia
foo(x::String, y::String) = println("My inputs x and y are both strings!")
```

Output:
```
foo (generic function with 1 method)
```

We see here that in order to restrict the type of `x` and `y` to `String`s, we just follow the input argument name by a double colon and the keyword `String`.

Now we'll see that `foo` works on `String`s and doesn't work on other input argument types.

```julia
foo("hello", "hi!")
```

Output:
```
My inputs x and y are both strings!
```

```julia
foo(3, 4)
```

Output:
```
Stacktrace:
 [1] top-level scope at In[6]:1
```

To get `foo` to work on integer (`Int`) inputs, let's tack `::Int` onto our input arguments when we declare `foo`.

```julia
foo(x::Int, y::Int) = println("My inputs x and y are both integers!")
```

Output:
```
foo (generic function with 2 methods)
```

```julia
foo(3, 4)
```

Output:
```
My inputs x and y are both integers!
```

Now `foo` works on integers! But look, `foo` also still works when `x` and `y` are strings!

```julia
foo("hello", "hi!")
```

Output:
```
My inputs x and y are both strings!
```

This is starting to get to the heart of multiple dispatch. When we declared

```julia
foo(x::Int, y::Int) = println("My inputs x and y are both integers!")
```
we didn't overwrite or replace
```julia
foo(y::String, y::String)
```
Instead, we just added an additional ***method*** to the ***generic function*** called `foo`.

A ***generic function*** is the abstract concept associated with a particular operation.

For example, the generic function `+` represents the concept of addition.

A ***method*** is a specific implementation of a generic function for *particular argument types*.

For example, `+` has methods that accept floating point numbers, integers, matrices, etc.

We can use the `methods` to see how many methods there are for `foo`.

```julia
methods(foo)
```

Output:
```
# 2 methods for generic function "foo":
[1] foo(x::Int64, y::Int64) in Main at In[7]:1
[2] foo(x::String, y::String) in Main at In[4]:1
```

Aside: how many methods do you think there are for addition?

```julia
methods(+)
```

Output:
```
# 166 methods for generic function "+":
[1] +(x::Bool, z::Complex{Bool}) in Base at complex.jl:286
[2] +(x::Bool, y::Bool) in Base at bool.jl:96
[3] +(x::Bool) in Base at bool.jl:93
[4] +(x::Bool, y::T) where T<:AbstractFloat in Base at bool.jl:104
[5] +(x::Bool, z::Complex) in Base at complex.jl:293
[6] +(a::Float16, b::Float16) in Base at float.jl:398
[7] +(x::Float32, y::Float32) in Base at float.jl:400
[8] +(x::Float64, y::Float64) in Base at float.jl:401
[9] +(z::Complex{Bool}, x::Bool) in Base at complex.jl:287
[10] +(z::Complex{Bool}, x::Real) in Base at complex.jl:301
[11] +(::Missing, ::Missing) in Base at missing.jl:114
[12] +(::Missing) in Base at missing.jl:100
[13] +(::Missing, ::Number) in Base at missing.jl:115
[14] +(level::Base.CoreLogging.LogLevel, inc::Integer) in Base.CoreLogging at logging.jl:107
[15] +(c::BigInt, x::BigFloat) in Base.MPFR at mpfr.jl:414
[16] +(a::BigInt, b::BigInt, c::BigInt, d::BigInt, e::BigInt) in Base.GMP at gmp.jl:515
[17] +(a::BigInt, b::Big
```

So, we now can call `foo` on integers or strings. When you call `foo` on a particular set of arguments, Julia will infer the types of the inputs and dispatch the appropriate method. *This* is multiple dispatch.

Multiple dispatch makes our code generic and fast. Our code can be generic and flexible because we can write code in terms of abstract operations such as addition and multiplication, rather than in terms of specific implementations. At the same time, our code runs quickly because Julia is able to call efficient methods for the relevant types.

To see which method is being dispatched when we call a generic function, we can use the @which macro:

```julia
@which foo(3, 4)
```

Output:
```
foo(x::Int64, y::Int64) in Main at In[7]:1
```

Let's see what happens when we use `@which` with the addition operator!

```julia
@which 3.0 + 3.0
```

Output:
```
+(x::Float64, y::Float64) in Base at float.jl:401
```

And we can continue to add other methods to our generic function `foo`. Let's add one that takes the ***abstract type*** `Number`, which includes subtypes such as `Int`, `Float64`, and other objects you would think of as numbers:

```julia
foo(x::Number, y::Number) = println("My inputs x and y are both numbers!")
```

Output:
```
foo (generic function with 3 methods)
```

This method for `foo` will work on, for example, floating point numbers:

```julia
foo(3.0, 4.0)
```

Output:
```
My inputs x and y are both numbers!
```

We can also add a fallback, duck-typed method for `foo` that takes inputs of any type:

```julia
foo(x, y) = println("I accept inputs of any type!")
```

Output:
```
foo (generic function with 4 methods)
```

Given the methods we've already written for `foo` so far, this method will be called whenever we pass non-numbers to `foo`:

```julia
v = rand(3)
foo(v, v)
```

Output:
```
I accept inputs of any type!
```

### Exercises

#### 9.1

Extend the function `foo`, adding a method that takes only one input argument, which is of type `Bool`, and returns "foo with one boolean!"

```julia
foo(x::Bool) = "foo with one boolean!"
```

Output:
```
foo (generic function with 5 methods)
```

#### 9.2

Check that the method being dispatched when you execute 
```julia
foo(true)
```
is the one you wrote.

```julia
foo(true)
```

Output:
```
"foo with one boolean!"
```

<!-- grade_id: cell-14072e60ae07c1a2, points: 1, grade: true, solution: false -->
```julia
@assert foo(true) == "foo with one boolean!"
```

Please click on `Validate` on the top, once you are done with the exercises.

```julia
# "hello" + "world!"
```

Output:
```
Stacktrace:
 [1] top-level scope at In[32]:1
```

```julia
add_strings(x::String, y::String) = string(x, y)
```

Output:
```
add_strings (generic function with 1 method)
```

```julia
add_strings("hi", "jey")
```

Output:
```
"hijey"
```

```julia

```